---
title: 'Session 4: Data Types'
output:
  learnr::tutorial:
    progressive: yes
    allow_skip: yes
runtime: shiny_prerendered
description: |
  Dive into the details of some special data types: factors, dates and times, and string and continue working on your Data Project.
---

```{r setup, include=FALSE}
# Ensure that libraries are loaded.
library(tidyverse)
# includes stringr:: for string manipulation 
# includes forcats:: for factor manipulation
# includes readr:: for (importing and) parsing data
library(learnr)
library(gradethis)
library(knitr)
library(kableExtra)
# New packages (must be installed before taking this tutorial)
library(lubridate) # date&time manipulation
library(broom) # presenting statistical results
library(tm) # to create & handle text corpora
library(wordcloud) # for word clouds
library(RColorBrewer) # color palettes (for word clouds; also loaded by wordcloud package)
library(topicmodels) # topic modelling

tutorial_options(exercise.timelimit = 60, exercise.checker = gradethis::grade_learnr)
knitr::opts_chunk$set(echo = FALSE, warning = FALSE, message = FALSE)
```

```{r, context="data", include=FALSE}
# Ensure that the data is loaded for the remainder of this tutorial.
Brexit <- UsingRTutorials::Brexit
surveyweek1 <- UsingRTutorials::surveyweek1
```

<!-- Define programming tip style -->
<style>
.tip {
  background-color: #f5f5f5;
}
</style>

<!-- Define question style -->
<style>
.question {
  color: #5A9DDB;
}
</style>

## Factors  

A factor is the R data type for a categorical (nominal or ordinal measurement level) variable. It may remind you of a variable with value labels in SPSS, but it is different. And it is a bit complicated.

The basic question is: When is it better to use a factor than a character variable?

### Character variables

R treats a character variable as a categorical variable (nominal measurement level).

Our example data set `Brexit` contains 600 articles about Brexit from the (former) website nujij.nl.

<div class="question" >
Create a frequency table of the variable `category` in the `Brexit`dataset.
</div>

```{r freqtable, exercise = TRUE}

```

<!-- To hide the solution, use a textual hint. -->
<div id="freqtable-hint">
__Hint:__ Use tidyverse!
</div>

```{r freqtable-solution}
Brexit %>% count(category)
```

```{r freqtable-check}
gradethis::grade_code(
  correct = "That wasn't hard, was it?", 
  incorrect = ""
  )
```

###

Note the order of values in this frequency table: alphabetical. 

In addition, note the negligible number of observations for media, etc.

### Character variable in statistical analysis

In statistical analysis, a character variable is usually converted into dummies. Regression models in R do this automatically.

Below are the results of a regression model predicting the number of votes for a post from the post category.

```{r regression, exercise = TRUE}
# Linear regression: more in Session 6.
lm(votes ~ category, data = Brexit) %>%
  broom::tidy() %>% #retrieving the main results
  kable(digits = c(0,2,2,2,3),
               col.names = c("Parameter", "b", "SE", "t", "p"),
               align = "lrrrr") %>%
  kable_styling() %>%
  kable_paper(full_width = FALSE)
```

```{r quiz}
question("What is the reference category here?",
  answer("(Intercept)"),
  answer("Algemeen", correct = TRUE),
  answer("Economie"),
  answer("Media"),
  answer("Overig"),
  answer("Politiek"),
  answer("Showbizz"),
  answer("Wetenschap")
)
```

Want another reference category or reorder the categories? Create a factor.

<div class="tip" >
__Programming Tip__

- Having the values in another order than alphabetically is the main reason for using a factor instead of a character variable.
</div>

### Create a factor   

<div class="question" >
Create a factor called `cat_fact` from the variable `category` and add it to tibble `Brexit`.
</div>

```{r factor, exercise = TRUE}

```

<!-- To hide the solution, use a textual hint. -->
<div id="factor-hint">
__Hint:__ Use the (base::) `factor()` function.
</div>

```{r factor-solution}
Brexit <- Brexit %>% mutate(cat_fact = factor(x = category))
```

```{r factor-check}
gradethis::grade_code(
  correct = "", 
  incorrect = ""
  )
```

###

Have a close look at the description of the new factor:

```{r factorinfo, exercise = TRUE, exercise.setup = "factor-solution", exercise.eval = TRUE}
# Info shown in Environment tab if you work in RStudio.
str(Brexit$cat_fact)
```

The factor contains a set of labels (levels) plus a set of integer values.

In contrast to SPSS, we must use the labels, we cannot use the factor integer values!

### Change the first factor category 

<div class="question" >
Add a factor called `cat_new` to tibble `Brexit` in which `Economie` (_Economy_) is the first category.
</div>

```{r relevel, exercise = TRUE, exercise.setup = "factor-solution"}

```

<!-- To hide the solution, use a textual hint. -->
<div id="relevel-hint">
__Hint:__ Use `fct_relevel()`, which is part of the `forcats` package, which is loaded by the `tidyverse` package.
</div>

```{r relevel-solution, exercise.setup = "factor-solution"}
Brexit <- Brexit %>% mutate(cat_new = fct_relevel(cat_fact, "Economie"))
```

```{r relevel-check}
gradethis::grade_code(
  correct = "", 
  incorrect = "There are more ways of getting the result but we want you to use the function you learned in the book."
  )
```

###

Now check that `Economie` is the reference category in a regression model.

```{r relevelmodel, exercise = TRUE, exercise.setup = "relevel-solution"}

```

<!-- To hide the solution, use a textual hint. -->
<div id="relevelmodel-hint">
__Hint:__ Adapt the code for a regression model from a previous exercise but do not use the `broom::tidy()`, `kable()`, and `kable_styling()`, and `kable_paper()` functions. Details on statistical analysis in R follow in Session 6.
</div>

```{r relevelmodel-solution}
lm(votes ~ cat_new, data = Brexit)
```

```{r relevelmodel-check}
gradethis::grade_code(
  correct = "", 
  incorrect = ""
  )
```

Note that we save the changes in the `Brexit` tibble because we want to use the reordered factor later on in this tutorial. But we use a new name for the new variable, so the original variable is preserved.

### Recode factor levels

<div class="question" >
Categories `Media`, `Overig` (_Other_), `Showbizz`, and `Wetenschap` (_Science_) contain only one observation. Merge them with category `Algemeen` (_General_) in a new variable `cat_fact2` in tibble `Brexit`. </div>

```{r recodefactor, exercise = TRUE, exercise.setup = "factor-solution"}

```

<!-- To hide the solution, use a textual hint. -->
<div id="recodefactor-hint">
__Hint:__ Use variable `cat_fact` as your starting point..
</div>

```{r recodefactor-solution}
Brexit <- Brexit %>% mutate(cat_fact2 = forcats::fct_recode(cat_fact, "Algemeen" = "Media", "Algemeen" = "Overig", "Algemeen" = "Showbizz", "Algemeen" = "Wetenschap"))
```

```{r recodefactor-check}
gradethis::grade_code(
  correct = "In the `fct_recode()`function, the new value is to the left, the old value is to the right. As in y <- x.", 
  incorrect = "Perhaps your order of factor values is different but your result can be OK. The code checker can only deal with one order. Poor soul that it is."
  )
```

### Grouping a numeric variable into a factor

Variable `reactions` in tibble `Brexit` gives the number of reactions that a post received. We would like to create a new variable that groups the number of reactions into three bins. Each bin should contain more or less the same number of observations (posts).

Unfortunately, variable `reactions` is not numeric.

<div class="question" >
Use function `parse_number()` to create a new numeric variable `react_num` in tibble `Brexit`, which gives us the number of reactions to a post.
</div>

```{r groupreactions1, exercise = TRUE}

```

<!-- To hide the solution, use a textual hint. -->
<div id="groupreactions1-hint">
__Hint:__ Function `parse_number()` is part of the `readr` package, which is loaded by the `tidyverse` package.
</div>

```{r groupreactions1-solution}
Brexit <- Brexit %>% mutate(react_num = parse_number(reactions))
```

```{r groupreactions1-check}
gradethis::grade_code(
  correct = "We use `parse_number()` to extract the (first) number from a string. Your first string manipulation in this tutorial! More to follow.", 
  incorrect = "Use tidyverse code with piping!"
  )
```

###

The three reactions bins should be labeled "low", "medium", "high". But the substantive order from low to high does not match the alphabetical order. The bins variable must therefore be a factor.

<div class="question" >
Group the number of reactions in three bins with (nearly) the same number of cases and store this variable as a factor named `react_cat` in tibble `Brexit`.
</div>

```{r orderedfactor, exercise = TRUE, exercise.setup = "groupreactions1-solution"}

```

<!-- To hide the solution, use a textual hint. -->
<div id="orderedfactor-hint">
__Hint:__ There are two steps: (1) bin the variable, (2) create a factor from the bins. The two steps can be performed in one `mutate()` function. It is even possible to use the same name for the new variable in both steps.
</div>

```{r orderedfactor-solution, exercise.setup = "groupreactions1-solution"}
Brexit <- Brexit %>% mutate(react_cat = ntile(react_num, 3), react_cat = factor(react_cat, levels = c(1, 2, 3), labels = c("low", "medium", "high")))
```

```{r orderedfactor-check}
gradethis::grade_code(
  correct = "", 
  incorrect = "Note the `levels` (sets the order) and `labels` (sets the label/value) arguments of the `factor()` function."
  )
```

###

Plot a bar chart of the factor representing the reactions bins. Are the bins in the desired order?

<div class="question" >
</div>

```{r binplot, exercise = TRUE, exercise.setup = "orderedfactor-solution"}

```

<!-- To hide the solution, use a textual hint. -->
<div id="binplot-hint">
__Hint:__ hinttext.
</div>

```{r binplot-solution}
ggplot(Brexit) + geom_bar(mapping = aes(x = react_cat))
```

```{r binplot-check}
gradethis::grade_code(
  correct = "Yes indeed, all three bins have the same number of observations (posts).", 
  incorrect = ""
  )
```

<div class="tip" >
__Programming Tip__

- An ordered factor assigns ordinal measurement level. In a regression model, no ordinary dummy variables are created for an ordered factor. So avoid ordered factors, unless you really need an ordinal variable.
</div>

## Dates and Times  

### Date-time data object

Example: Publication date of post on _nujij.nl_ in tibble `Brexit`.

```{r POSIXct, echo=FALSE}
str(Brexit$pubDate, "%B %d, %Y %H:%M:%S")
```

Date plus time of day are (usually) stored as a `POSIXct` data type.

- Number of seconds since the beginning of 1970.
- Shown in the standard format (of your computer locale).

### String to date-time  

In data set `Brexit`, variable `pubDate` specifies the date and time an article was posted as a string, e.g., "`r Brexit$pubDate[1]`". 

<div class="question" >
Create a date-time variable `pubDateTime` in tibble `Brexit` using a parse function in package `readr` (p. 134 {Ch. 11.3.4}).
</div>

```{r pubDate1, exercise = TRUE}

```

<!-- To hide the solution, use a textual hint. -->
<div id="pubDate1-hint">
__Hint:__ Carefully specify the date and time formatting codes. See help for `parse_datetime()` under __Format specification__.
</div>

```{r pubDate1-solution}
Brexit <- Brexit %>% mutate(pubDateTime = parse_datetime(pubDate, "%B %d, %Y %H:%M:%S"))
```

```{r pubDate1-check}
gradethis::grade_code(
  correct = "Compare the new variable to the original publication date variable: Is the new variable correct?", 
  incorrect = ""
  )
```

###

<div class="question" >
Now do the same with the `lubridate` package (p. 239 {Ch. 16.2.1}), storing the results as a new variable `pubDateTime2` in tibble `Brexit`.
</div>

```{r pubDate2, exercise = TRUE}

```

<!-- To hide the solution, use a textual hint. -->
<div id="pubDate2-hint">
__Hint:__ Select the correct function. Check out help for `ymd_hms()`.
</div>

```{r pubDate2-solution}
Brexit <- Brexit %>% mutate(pubDateTime2 = mdy_hms(pubDate))
```

```{r pubDate2-check}
gradethis::grade_code(
  correct = "Again, compare the new variable to the original publication date variable: Is the new variable correct?", 
  incorrect = ""
  )
```

- Package `lubridate` simplifies the identification of date/time elements.
- Date and time formatting codes are also used for displaying dates and times.

<div class="tip" >
__Programming Tip__

- Especially when you let `readr` guess the date and time format, the results can be wrong. Always compare the new date-time variable to the original string variable.
</div>

### Retrieve info from date-time

It is easy to get the date from a date-time (POSIXct) variable with the `lubridate` package.

<div class="question" >
Add a variable `date` to the `Brexit` tibble, using the recently created variable `pubDateTime` as your starting point.
</div>

```{r Date, exercise = TRUE, exercise.setup = "pubDate1-solution"}

```

<!-- To hide the solution, use a textual hint. -->
<div id="Date-hint">
__Hint:__ Don't forget to pipe.
</div>

```{r Date-solution, exercise.setup = "pubDate1-solution"}
Brexit <- Brexit %>% mutate(date = date(pubDateTime))
```

```{r Date-check}
gradethis::grade_code(
  correct = "", 
  incorrect = ""
  )
```

###

<div class="question" >
Inspect the dates of the posts with a frequency polygon.
</div>

```{r datePlot, exercise = TRUE, exercise.setup = "Date-solution"}

```

<!-- To hide the solution, use a textual hint. -->
<div id="datePlot-hint">
__Hint:__ Pay attention to the `stat` argument.
</div>

```{r datePlot-solution}
Brexit %>% ggplot(aes(x = date)) + geom_freqpoly(stat = "count")
```

```{r datePlot-check}
gradethis::grade_code(
  correct = "", 
  incorrect = "The aesthetics for the horizontal axis can be specified in the `ggplot()` function or in the `geom_freqpoly()` geom."
  )
```

###

- Note that there is a `date` variable type in R. It behaves just like a `date-time` variable.
- `date()` is also a base R function. When you load the `lubridate` package, the base R `date()` function is overriden.

###

It is easy to get the weekday with the `lubridate` package.

<div class="question" >
Store the weekday as a new variable `weekday` containing the names of the weekdays in the tibble `Brexit`.
</div>

```{r weekday, exercise = TRUE, exercise.setup = "pubDate1-solution"}

```

<!-- To hide the solution, use a textual hint. -->
<div id="weekday-hint">
__Hint:__ How about the `wday()` function?
</div>

```{r weekday-solution}
Brexit <- Brexit %>% mutate(weekday = wday(pubDateTime, label=TRUE))
```

```{r weekday-check}
gradethis::grade_code(
  correct = "", 
  incorrect = ""
  )
```

###

```{r plotweekday, echo=FALSE}
Brexit  %>% mutate(weekday = wday(parse_datetime(pubDate, "%B %d, %Y %H:%M:%S"), label=TRUE)) %>% ggplot(aes(weekday)) + geom_bar()
```

Variable `weekday` must be a factor. How can you tell?

###  

It is trickier to get the time of day because R has no time data type. The book offers a trick.

<div class="question" >
Apply this trick to the `pubDateTime` variable in tibble `Brexit` to create a variable `time`. Show the variable as an area plot; don't save it.
</div>

```{r time, exercise = TRUE, exercise.setup = "pubDate1-solution"}

```

<!-- To hide the solution, use a textual hint. -->
<div id="time-hint">
__Hint:__ See Section _Setting Components_ in the book.
</div>

```{r time-solution}
Brexit %>% mutate(time = update(pubDateTime, yday = 1)) %>% ggplot(aes(time)) + geom_area(stat = "bin", bins = 30) + scale_x_datetime(date_labels = "%H")
```

```{r time-check}
gradethis::grade_code(
  correct = "", 
  incorrect = ""
  )
```

### Calculating with date-time

Let us have a look at the time span between successive posts. How long does it take for a new post to appear?

<div class="question" >
Subtract the `pubDateTime` of the previous post from the `pubDateTime` of the current post and store the result as variable `lapse` in tibble `Brexit`.
</div>

```{r timelapse, exercise = TRUE, exercise.setup = "pubDate1-solution"}

```

<!-- To hide the solution, use a textual hint. -->
<div id="timelapse-hint">
__Hint:__ Sort the data to be sure that they are in chronological order.
</div>

```{r timelapse-solution, exercise.setup = "pubDate1-solution"}
Brexit <- Brexit %>% arrange(pubDateTime) %>% mutate(lapse = pubDateTime - lag(pubDateTime, n = 1L))
```

```{r timelapse-check}
gradethis::grade_code(
  correct = "", 
  incorrect = ""
  )
```

###

```{r timelapseplot, echo=FALSE}
#Show publication date-time and lapse for first few observations.
Brexit %>% 
  mutate(pubDateTime = parse_datetime(pubDate, "%B %d, %Y %H:%M:%S")) %>%
  arrange(pubDateTime) %>% mutate(lapse = pubDateTime - lag(pubDateTime, n = 1L)) %>% 
  select(pubDateTime, lapse) %>%
  slice(1:6)
```

The time difference between subsequent postings are clearly in seconds (see above plot). What if we use the dates of publications?

<div class="question" >
Show the difference in days between the first six subsequent posts (new variable `lapse2`) as in the above table.
</div>

```{r daylapse, exercise = TRUE, exercise.setup = "Date-solution"}

```

<!-- To hide the solution, use a textual hint. -->
<div id="daylapse-hint">
__Hint:__ Don't save the new variable, just show it in a table.
</div>

```{r daylapse-solution}
Brexit %>% arrange(date) %>% mutate(lapse2 = date - lag(date, n = 1L)) %>% select(date, lapse2) %>% slice(1:6)
```

```{r daylapse-check}
gradethis::grade_code(
  correct = "", 
  incorrect = ""
  )
```

###

Now the difference is expressed in days. Subtracting two date-times yields data type `difftime`, which can have several time units. It is the largest time unit that is smaller than the smallest measured time difference.

If you want to be sure that second is the time unit of the date-time difference, change the `difftime` into a `duration` with the `as.duration()` function.

<div class="question" >
Plot variable `lapse` as a histogram with 50 bins.
</div>

```{r plotlapse, exercise = TRUE, exercise.setup = "timelapse-solution"}

```

<!-- To hide the solution, use a textual hint. -->
<div id="plotlapse-hint">
__Hint:__ Use `as.numeric()` to tell ggplot that the `difftime` data object must be treated as a numerica variable.
</div>

```{r plotlapse-solution}
Brexit %>% ggplot(aes(x = as.numeric(lapse))) + geom_histogram(bins = 50)
```

```{r plotlapse-check}
gradethis::grade_code(
  correct = "", 
  incorrect = ""
  )
```

<div class="tip" >
__Programming Tips__

- Work with `duration` instead of `difftime` unless you have to take into account daylight saving and leap seconds.
- Note: For minutes, hours, days, and weeks, a `duration` can be divided (ignoring leap seconds) by the appropriate number of seconds.
</div>

## Strings  

### String manipulation  

Main applications:  
  
1. Change string into another type (factor, date, number) that is more convenient for analysis.  
    - Examples in previous topics of this tutorial (string to factor or date-time).  
2. Query and quantify natural text contents.  
    - Find texts with particular contents.  
    - Change content characteristics into variables.  
    - Not applicable to all complex data sets.

Topic 2 is esp. relevant for students working on a complex data set including text (natural language): nujij.nl, German Chancellor Debate (debate statement text available in variable `TEXT`).

### Change variable names

Friends & Families weekly survey: nasty variable names. Use `rename()`? 

```{r imagesurvey, out.width="100%"}
knitr::include_graphics("images/surveyweek1.png")
```
 
- Data sets in the wild often contain troublesome variable names.
- We can use `rename()` to change each name.
- Instead, let's handle variable names as string data and use string manipulation techniques to simplify the variable names.

###

We want to simplify the variable name to the date for which the sad/depressed score was entered.

<div class="tip" >
__Programming Tip__

- If possible, first create code that works for one example. Next, extend the code to deal with all cases.
</div>

Let us first deal with the first troublesome variable name in tibble `surveyweek1`, which is the fourth variable. 

```{r getvar, exercise=TRUE}
# Get a variable name as a character variable with the names() function.
names(surveyweek1)[4]
# [4] selects item 4 from the vector of variable names
```

### Manipulate one variable name

Step 1: Remove all up to (and including) the first " - ".

<div class="question" >
Run the below code and inspect the intermediary results to check that you understand what each line of code does. For example, why is there `+ 1` within the `str_sub()` function?
</div>

```{r step1, exercise = TRUE, exercise.lines = 14}
# Store variable name in a data object, so we can use it.
varname <- names(surveyweek1)[4]
# Show result to check that the code works.
varname

# Find the first position of " - ".
# Show results for understanding before using it below.
str_locate(varname, " - ") #start and end position!

# Select everything after the end position of " - ".
varname <- str_sub(varname, start = str_locate(varname, " - ")[2] + 1)

# Show result to check that the code works.
varname
```
###

The book focuses on finding and replacing partcular strings/words:

- `str_locate()` with `str_sub()` nice for cutting up a string.
- `[2]` behind str_locate() selects the second item of the vector: the end position.

### 

Step 2: Get rid of the weekday name, so remove all up to (and including)  ", ".

```{r step2, exercise = TRUE, exercise.setup = "step1", exercise.lines = 6}
# Show the name that we start with.
varname
# Select substring after the end position of ", ".
varname <- str_sub(varname, start = str_locate(varname, ", ")[2] + 1)
# Show result: see the difference.
varname
```
###

Step 3: Get rid of the "Sad or depressed" part, so remove all from (and including)  " - ".

```{r step3, exercise = TRUE, exercise.setup = "step2", exercise.lines = 6}
# Show the name that we start with.
varname
# Select substring up to start position of " - ".
varname <- str_sub(varname, end = str_locate(varname, " - ")[1] - 1)
# Show result: see the difference.
varname
```

<div class="tip" >
__Programming Tip__

- It is good practice to check intermediary results, so you know that your code performs as expected.
- It is equally good practice to comment out code that is only used for checking results when the code is definitive.
</div>

###

Note that `start` and `end` arguments can be combined in `sub_str()`. Steps 2 and 3 can thus be done in one command.

<div class="question" >
Write the `str_sub()`function that combines Steps 2 and 3. 
</div>

```{r step2and3, exercise = TRUE, exercise.setup = "step1"}

```

<!-- To hide the solution, use a textual hint. -->
<div id="step2and3-hint">
__Hint:__ Don't store the result in `varname`; just use the `str_sub()` function.
</div>

```{r step2and3-solution}
str_sub(varname, start = str_locate(varname, ", ")[2] + 1, end = str_locate(varname, " - ")[1] - 1)
```

```{r step2and3-check}
gradethis::grade_code(
  correct = "You added the `end` argument at the correct position.", 
  incorrect = ""
  )
```

### Manipulate all variable names

Now, let us apply the commands to all relevant variables (sequence 4 to 10).

<div class="question" >
Run the below code. What do `names()` and `str_locate()` produce now?
</div>

```{r strlocate1, exercise = TRUE, exercise.lines = 7}
# Get all selected variable names.
varname <- names(surveyweek1)[4:10] 
# Check result.
varname

# Get the start and end positions.
str_locate(varname, " - ")
```

###

Now, str_locate creates a matrix with start and end position columns for each variable. If we want the start positions or the end positions, we must select the first column `[,1]` or the second column `[,2]` from the matrix.

```{r strlocate2, exercise = TRUE, exercise.setup = "strlocate1"}
# Get the end positions from the matrix.
str_locate(varname, " - ")[,2]
```

### 

<div class="question" >
Adjust the commands that we have created for changing one variable name to changing the names of all relevant variables.
</div>

```{r changeall, exercise = TRUE}
# Store variable name in a data object, so we can use it.
varname <- names(surveyweek1)[4]
# Select everything after the end position of " - ".
varname <- str_sub(varname, start = str_locate(varname, " - ")[2] + 1)
# Select substring after the end position of ", ".
varname <- str_sub(varname, start = str_locate(varname, ", ")[2] + 1)
# Select substring up to start position of " - ".
varname <- str_sub(varname, end = str_locate(varname, " - ")[1] - 1)
# Check result.
varname
```

```{r changeall-check}
gradethis::grade_result(
  fail_if(~ identical(.result, "March 8"), "Did you select variables 4 to 10 in the first command?"),
  fail_if(~ length(.result) == 7 && .result[2] == " March ", "Did you use columns instead of single values for start and end positions?"),
  pass_if(~ identical(.result, c("March 8", "March 9", "March 10", "March 11", "March 12", "March 13", "March 14")), "You have selected all seven variables (`[4:10]`) and the correct columns of start (`[,1]` instead of `[1]`) and end positions (`[,2]` instead of `[2]`).")
)
```

###

Once we have the correct variable names, we can use them to replace the names in the `surveyweek1` tibble.

```{r replacenames-setup}
# Store variable name in a data object, so we can use it.
varname <- names(surveyweek1)[4:10]
# Select everything after the end position of " - ".
varname <- str_sub(varname, start = str_locate(varname, " - ")[,2] + 1)
# Select substring after the end position of ", ".
varname <- str_sub(varname, start = str_locate(varname, ", ")[,2] + 1)
# Select substring up to start position of " - ".
varname <- str_sub(varname, end = str_locate(varname, " - ")[,1] - 1)
```

```{r replacenames, exercise = TRUE}
# Replace variable names in tibble.
names(surveyweek1)[4:10] <- varname
# Inspect results
str(surveyweek1)
```

### 

Because the variable names are so well structured in this particular example, we can achieve the same with the tidyverse `separate()` function.

<div class="question" >
Inspect the results of each step in the pipe with `View()`. Note that the View window may be hidden behind your RStudio screen.
</div>

```{r separate, exercise = TRUE, exercise.lines = 12}
#tidyverse requires a data frame
varnames <- data.frame(name = names(surveyweek1)[4:10]) %>%
  #split each name into three parts
  separate(name, into = c("humbug", "date", "topic"), sep = " - ") %>%
  #split the date into weekday and date
  separate(date, into = c("weekday", "date"), sep = ", ") %>%
  #extract the date variable as a vector
  pull(date)
# Replace variable names in tibble.
names(surveyweek1)[4:10] <- varnames
# Show result.
str(surveyweek1)
```

- Tidyverse functions only operate on data frames but the variable names are a vector: just a series of names.
- Therefore we must change them into a data frame at the start with `data.frame()` and transform them back to a vector (without variable name) at the end with `pull()`.

## Fun Stuff

Natural text such as the header or body text of posts is used for classifying texts (topic modelling), for example, using sentiment analysis, or investigating relations between words, e.g., as visualized in a word cloud. 

### Word cloud

The special R data class `Corpus` has been developed to handle sets of texts. The `tm` package offers nice tools for creating, cleaning, and manipulating a corpus of texts. Let us use this package and the `wordcloud` and `wordcloud2` packages to create some word clouds.

<div class="question" >
Use help and inspect intermediary results to understand what happens here.
</div>

```{r wordcloud1, exercise = TRUE, exercise.lines = 23, exercise.eval = TRUE}
# Load the required packages.
library(tm)
library(wordcloud)
# Create a corpus for the headings (titles) of the Brexit articles.
docs <- Brexit %>%
  # extract the titles as a vector
  pull(title) %>%
  # use tm functions to create a corpus
  VectorSource() %>%
  Corpus() %>%
  # use tm functions to clean the data
  tm_map(removeNumbers) %>% #remove numbers (assuming they are not informative)
  tm_map(removePunctuation) %>% #remove punctuation (assuming they are not informative)
  tm_map(stripWhitespace) %>% #remove blanks
  tm_map(content_transformer(tolower)) %>% #only lower case letters
  tm_map(removeWords, stopwords("dutch")) #remove Duch stopwords
# Create our first wordcloud.
wordcloud(
  docs, 
  random.order = FALSE,
  colors = brewer.pal(11, "Spectral")
  )
```

###

The word _Brexit_ clearly dominates the word cloud. Well, articles have been selected because this word appears. Let us make a word cloud without it.

<div class="question" >
Play around with the selection of words and the arguments of the `wordcloud()` function.
</div>

```{r exerciselabel, exercise = TRUE, exercise.eval = TRUE}
docs %>% 
  #remove word "brexit"
  tm_map(removeWords, c("brexit")) %>%
  #create wordcloud
  wordcloud(
    random.order = FALSE,
    colors = brewer.pal(11, "Spectral")
  )
```

### Topic modelling

Usually, a document-term-matrix is required for classifying texts based on the (co-)occurrence of words. A document-term-matrix is a data matrix with rows (or columns) for texts (documents) and a column (or row) for each word that appears in at least one text. The cells of the matrix contain frequencies: How many times does this word occur in this document. The `tm` package can create document-term-matrices from a text corpus. 

```{r topicmodel, exercise = TRUE, exercise.lines = 25, exercise.eval = TRUE}
# Load the required package.
library(topicmodels)
# Create a document-term-matrix from the corpus.
dtm = DocumentTermMatrix(docs)
# Run a Latent Dirichlet Allocation topic model with 2 topics
lda = LDA(dtm, k = 2, seed = 234) 
# Calculate the probabilities that a text or word belongs to Topic 1 or Topic 2.
pred = posterior(lda,dtm) 

# Show word clouds of terms that are more strongly related to one topic than to the other.
# First, create a tibble with the ratio of the probability that a term links with Topic 2 over the probability that a term links with Topic 1.
probRatio <- pred$terms %>% as_tibble() %>%
  rownames_to_column(var = "topic") %>% #preserve topic number as variable
  #turn term variables into one variable
  pivot_longer(!topic, names_to = "term", values_to = "prob") %>% 
  #calculate the probability ration
  group_by(term) %>%
  arrange(term, topic) %>%
  mutate(pratio = prob / lag(prob, n = 1L)) %>%
  #only retain the rows with a valid probability ratio (namely, topic 2)
  filter(topic == 2)

# SHow two plots side by side.
par(mar = c(0, 0, 0, 0), mfcol = c(1, 2))
# Wordcloud Topic 1: terms that are at least 5 times more likely for topic 1.
nottopic1 <- probRatio %>%
  filter(pratio > 0.2) %>%
  pull(term) #we need a vector for the removeWords() function
docs %>% 
  #remove word "brexit"
  tm_map(removeWords, nottopic1) %>%
  #create wordcloud
  wordcloud(
    random.order = FALSE,
    colors = brewer.pal(11, "Spectral"),
    min.freq = 2
  )
# Wordcloud Topic 2: terms that are at least 5 times more likely for topic 1.
nottopic2 <- probRatio %>%
  filter(pratio < 5) %>%
  pull(term) #we need a vector for the removeWords() function
docs %>% 
  #remove word "brexit"
  tm_map(removeWords, nottopic2) %>%
  #create wordcloud
  wordcloud(
    random.order = FALSE,
    colors = brewer.pal(11, "Spectral"),
    min.freq = 2
  )
```

Well, what's the difference between the two topics? Note that the results may change dramatically from one run to the next. Perhaps, there is just one story.

## Data Project

* Change an interesting character variable in your complex data set into a factor.
* Reorder the groups (values) within the factor.
* Group a numeric variable into a factor with nice group labels.

Select a time or date variable in your complex data set.

* If it is not in date-time (or Date) format, change it into that format.
* Compare some values of the date-time variable to the values (strings) in the original .csv file. Have the values been correctly read?
* Graph the time or date variable to check the distribution.

<div class="tip" >
__Programming Tip__

- Never assume that dates and times are correctly read.
</div>
